```python
%matplotlib inline
%precision 16
from __future__ import print_function
import numpy
import matplotlib.pyplot as plt
```

Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel $\rightarrow$ Restart) and then run all cells (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says YOUR CODE HERE or "YOUR ANSWER HERE", as well as your name and collaborators below:

# Homework 0:  Introduction

## Question 1

Compute the solution to the following sets of equations.  **FULLY** justify your solution (do not just write the answer).

If you want to review some of these concepts check out Strang's [Linear Algebra](https://clio.columbia.edu/catalog/10612929) text or for only the fundamentals check out Strang's [The Fundamental Theorem of Linear Algebra](http://www.jstor.org/stable/2324660?seq=1#page_scan_tab_contents)

**(a)** (5) Solve $A x = b$ where
$$
    A = \begin{bmatrix}
        2 & 1 \\
        1 & 3
    \end{bmatrix} ~~~~ 
    b = \begin{bmatrix}
        1 \\
        -2
    \end{bmatrix}
$$

<!-- grade_id: A1-a, points: 5, grade: true, solution: true -->
#Renjie Li, UNI = RL2932

**Solution:**

I) With the concept of elimination, the above linear equation can be treated as two equations with two unknowns 
($x_{1}$and $x_{2}$are two components of vector $\bf{x}$):
  \begin{align}
  2x_{1}+x_{2}&=1  \\ 
  x_{1}+3x_{2}&=-2  
  \end{align}
   Substract 1/2 times equation 1 from equation 2:
  \begin{align}
  2x_{1}+x_{2}&=1 \\
  2.5x_{2}&=-2.5
  \end{align}
  Solution:
  $$
  x_{2}=\frac{-2.5}{2.5} = -1 \\
  x_{1}=\frac{1-x_{2}}{2}= 1
  $$
  Thus:
  $$
  x = \begin{bmatrix}
  x_{1} \\
  x_{2}
  \end{bmatrix}
  = \begin{bmatrix}
  1 \\
  -1
  \end{bmatrix}
  $$
  



II) With the concept of inverse matrices, we need to find the inverse matrix of $A$ :
  
  Since the determinant of $A$ is $ 2\cdot 3-1\cdot 1=5$ , $A$ matrix is invertible .
  And its inverse matrix is :
  $$
  A^{-1}=\begin{bmatrix}
  2 & 1\\
  1 & 3
  \end{bmatrix}^{-1}
  = \frac{1}{5}\begin{bmatrix}
  3 & -1\\
  -1 & 2
  \end{bmatrix} 
  = \begin{bmatrix}
  \frac{3}{5} & -\frac{1}{5} \\
  -\frac{1}{5} & \frac{2}{5}
  \end{bmatrix}
  $$
  Thus:
  $$
  x = A^{-1}Ax = A^{-1}b \\
  x =\begin{bmatrix}
  \frac{3}{5} & -\frac{1}{5} \\
  -\frac{1}{5} & \frac{2}{5}
  \end{bmatrix}
  \begin{bmatrix}
  1 \\
  -2
  \end{bmatrix}
  =\begin{bmatrix}
  1 \\
  -1
  \end{bmatrix}
  $$
  
  
  

**(b)** (5) Solve the system of equations:
\begin{align}
    2x + 3y &= 1 \\
    6x + 9y &= 3
\end{align}

<!-- grade_id: A1-b, points: 5, grade: true, solution: true -->
**Solution:**

With the concept of elimination, substract 3 times equation 1 from equation 2 :
$$
0y=0
$$
Every $y$ here satisfies the above equation. 
In fact there is only one equation, since the two equations in question description represent the same line.
We can choose the value of $y$ freely and then determine x by $x = \frac{1-3y}{2} $

**(c)** (5) Why will I not be able to solve $Ax=b$ for
$$
    A = \begin{bmatrix}
        2 & 1 \\
        1 & 3 \\
        0 & 4
    \end{bmatrix} ~~~~ 
    b = \begin{bmatrix}
        2 \\
        0 \\
        1
    \end{bmatrix}?
$$
Describe why this is a problem in terms of the column, row, left-null, and null spaces of $A$ and how that relates to the vector $b$.

<!-- grade_id: A1-c, points: 5, grade: true, solution: true -->
Description:

$Ax=b$ is solvable only if $b$ is in the column space of $A$.
$A$ is a 3 by 2 matrix, 2 columns cannot fill 3-dimensional vector space , which means some $b$ are not in the column space of $A$.

The row echelon form of $A$ is :
$ \begin{bmatrix}
2 & 1\\
0 & 2.5\\
0 & 0
\end{bmatrix}
$

It has two pivots, which means bo,th two columns define the column space of $A$.Thus the column space of $A$ is :
$
 x_{1} \begin{bmatrix}
 2 \\ 1 \\ 0
 \end{bmatrix}
 + x_{2}\begin{bmatrix}
 1 \\ 3 \\ 4
 \end{bmatrix}
$

The given vector $b$ here is not in the column space of $A$. That is why this equation is not solvable.


By the way, we can also use elimination on augmented matrix:
$
\begin {bmatrix}
2 & 1 & b_{1}\\
1 & 3 & b_{2}\\
0 & 4 & b_{3}
\end{bmatrix}
\rightarrow
\begin {bmatrix}
2 & 1 & b_{1} \\
0 & 2.5 & b_{2}-0.5b_{1} \\
0 & 0 & b_{3}-1.6b_{2}+0.8b_{1} 
\end{bmatrix}
$

$Ax=b$ is solvable when every entry in last row equals zero. In this case, $1-1.6\cdot 0+0.8\cdot 2 = 2.6 \neq 0 $ .
We can not solve this equation.

## Question 2

**(a)** (10) Write a function that computes
$$
    \sum^\infty_{n=1} \frac{a^n}{b^{n-1}}
$$
until the difference between subsequent partial sums is less than the given tolerance $T$.  Return the computed sum.  Make sure to include a way for the function to exit if the partial sums do not satisfy the above criteria (the sum may not be convergent for instance).

<!-- grade_id: A2, points: 0, grade: false, solution: true -->
```python
#
# YOUR CODE HERE
def compute_sum (a,b,T):
    sum = 0
    n = 1
    while ((a ** n) / (b**(n-1)) > T) :
        sum += (a ** n) / (b**(n-1))
    return sum
    
print (compute_sum(1.0, 2.0, 1e-16))
numpy.testing.assert_allclose(compute_sum(1.0, 2.0, 1e-16), -2.0)
try:
    compute_sum(1.0, 2.0, 1e-16)
except ValueError:
    pass
else:
    assert False

```

<!-- grade_id: T2-a, points: 10, grade: true, solution: false -->
```python
import numpy
numpy.testing.assert_allclose(compute_sum(-2.0, 3.0, 1e-16), -1.2)
numpy.testing.assert_allclose(compute_sum(1.0, 2.0, 1e-16), 2.0)
try:
    compute_sum(2.0, 1.0, 1e-16)
except ValueError:
    pass
else:
    assert False
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-2-5813cf8925ad> in <module>()
      1 import numpy
----> 2 numpy.testing.assert_allclose(compute_sum(-2.0, 3.0, 1e-16), -1.2)
      3 numpy.testing.assert_allclose(compute_sum(1.0, 2.0, 1e-16), 2.0)
      4 try:
      5     compute_sum(2.0, 1.0, 1e-16)

NameError: name 'compute_sum' is not defined
```

**(b)** (5) Explore different tolerances for your function above for $a=-2$ and $b=3$.  Plot the difference
$$
    \left| ~ \sum^\infty_{n=1} \frac{a^n}{b^{n-1}} - (-1.2) ~\right |
$$
versus the value of the tolerance.  The plotting command `loglog` may be useful to effectively visualize the problem.  What do you observe?

Hint:  Try using tolerances in the range $T \in [10^{-30}, 10^{-2}]$.

<!-- grade_id: A2-b1, points: 4, grade: true, solution: true -->
```python
#
# YOUR CODE HERE


```

<!-- grade_id: A2-b2, points: 1, grade: true, solution: true -->
#
# YOUR CODE HERE



## Question 3

**(a)** (7) Compute the first 3 terms of the Taylor series of the function
$$
    f(x) = e^{-x^2} \sin( x - \pi)
$$
centered at $x_0 = \pi$.

<!-- grade_id: A3-a, points: 7, grade: true, solution: true -->
Derivatives:
$$\begin{aligned}
f'(x) &= e^{-x^{2}}(2xsin(x)-cos(x)) \\
f''(x) &= e^{-x^{2}}[(3-4x^{2})sin(x)+4xcos(x)]
\end{aligned}
$$

Taylor polynomials:
$$T_N(x) = \sum^N_{n=0} \frac{f^{(n)}(x_0)\cdot(x-x_0)^n}{n!} $$
For centered at $x_{0}=\pi $
$$T_N(x) = \sum^N_{n=0} \frac{f^{(n)}(\pi)\cdot(x-\pi)^n}{n!}\Rightarrow $$

$$
\begin{aligned}
T_2(x) & = e^{-\pi^{2}}sin(\pi - \pi)+ e^{-\pi^{2}}[2\pi sin(\pi)-cos(\pi)](x-\pi)+\frac{1}{2}e^{-\pi^{2}}[(3-4\pi^{2})sin(\pi)+4\pi cos(\pi)](x-\pi)^{2}  \\
& = \pi e^{-\pi^{2}}(x-\pi)-2e^{-\pi^{2}}(x-\pi)^{2}
\end{aligned}
$$

**(b)** (8) Solve the ODE
$$
    u'' + u' + \frac{5}{4} u = 0
$$
with initial conditions $u(0) = 3$ and $u'(0) = 1$.  Plot the solution and comment on its behavior as $t \rightarrow \infty$.

<!-- grade_id: A3-b1, points: 6, grade: true, solution: true -->
#
# YOUR CODE HERE



<!-- grade_id: A3-b2, points: 2, grade: true, solution: true -->
```python
#
# YOUR CODE HERE


```

## Question 4

(5) Plot something fun making sure to label the axes and colorbar if appropriate.  Use the [matplotlib gallery](http://matplotlib.org/gallery.html) for inspiration.

<!-- grade_id: A4, points: 5, grade: true, solution: true -->
```python
#
# YOUR CODE HERE


```