```python
%matplotlib inline
%precision 16
from __future__ import print_function
import numpy
import matplotlib.pyplot as plt
```

Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel $\rightarrow$ Restart) and then run all cells (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says YOUR CODE HERE or "YOUR ANSWER HERE", as well as your name and collaborators below:

# Homework 0:  Introduction

## Question 1

Compute the solution to the following sets of equations.  **FULLY** justify your solution (do not just write the answer).

If you want to review some of these concepts check out Strang's [Linear Algebra](https://clio.columbia.edu/catalog/10612929) text or for only the fundamentals check out Strang's [The Fundamental Theorem of Linear Algebra](http://www.jstor.org/stable/2324660?seq=1#page_scan_tab_contents)

**(a)** (5) Solve $A x = b$ where
$$
    A = \begin{bmatrix}
        2 & 1 \\
        1 & 3
    \end{bmatrix} ~~~~ 
    b = \begin{bmatrix}
        1 \\
        -2
    \end{bmatrix}
$$

<!-- grade_id: A1-a, points: 5, grade: true, solution: true -->
```python
#
# YOUR CODE HERE
#Renjie Li, UNI = RL2932
import numpy
from numpy.linalg import inv
a = numpy.array([[2,1],[1,3]])
b = numpy.array([1,-2])
#x = numpy.linalg.solve(a, b)
#print (x)
inv_a = inv(a)
numpy.dot(inv_a,b)
```

Output:
```
array([ 1., -1.])
```

**(b)** (5) Solve the system of equations:
\begin{align}
    2x + 3y &= 1 \\
    6x + 9y &= 3
\end{align}

<!-- grade_id: A1-b, points: 5, grade: true, solution: true -->
```python
#
# YOUR CODE HERE
# Renjie Li , UNI = RL2932
a = numpy.array([[2,3],[6,9]])
b = numpy.array([1,3])
x = numpy.linalg.solve(a, b)
print (x)
np.allclose(np.dot(a, x), b)
```

Output:
```
[  2.77555756e-17   3.33333333e-01]
True
```

**(c)** (5) Why will I not be able to solve $Ax=b$ for
$$
    A = \begin{bmatrix}
        2 & 1 \\
        1 & 3 \\
        0 & 4
    \end{bmatrix} ~~~~ 
    b = \begin{bmatrix}
        2 \\
        0 \\
        1
    \end{bmatrix}?
$$
Describe why this is a problem in terms of the column, row, left-null, and null spaces of $A$ and how that relates to the vector $b$.

<!-- grade_id: A1-c, points: 5, grade: true, solution: true -->
#
# YOUR CODE HERE



## Question 2

**(a)** (10) Write a function that computes
$$
    \sum^\infty_{n=1} \frac{a^n}{b^{n-1}}
$$
until the difference between subsequent partial sums is less than the given tolerance $T$.  Return the computed sum.  Make sure to include a way for the function to exit if the partial sums do not satisfy the above criteria (the sum may not be convergent for instance).

<!-- grade_id: A2, points: 0, grade: false, solution: true -->
```python
#
# YOUR CODE HERE


```

<!-- grade_id: T2-a, points: 10, grade: true, solution: false -->
```python
numpy.testing.assert_allclose(compute_sum(-2.0, 3.0, 1e-16), -1.2)
numpy.testing.assert_allclose(compute_sum(1.0, 2.0, 1e-16), 2.0)
try:
    compute_sum(2.0, 1.0, 1e-16)
except ValueError:
    pass
else:
    assert False
```

**(b)** (5) Explore different tolerances for your function above for $a=-2$ and $b=3$.  Plot the difference
$$
    \left| ~ \sum^\infty_{n=1} \frac{a^n}{b^{n-1}} - (-1.2) ~\right |
$$
versus the value of the tolerance.  The plotting command `loglog` may be useful to effectively visualize the problem.  What do you observe?

Hint:  Try using tolerances in the range $T \in [10^{-30}, 10^{-2}]$.

<!-- grade_id: A2-b1, points: 4, grade: true, solution: true -->
```python
#
# YOUR CODE HERE


```

<!-- grade_id: A2-b2, points: 1, grade: true, solution: true -->
#
# YOUR CODE HERE



## Question 3

**(a)** (7) Compute the first 3 terms of the Taylor series of the function
$$
    f(x) = e^{-x^2} \sin( x - \pi)
$$
centered at $x_0 = \pi$.

<!-- grade_id: A3-a, points: 7, grade: true, solution: true -->
#
# YOUR CODE HERE



**(b)** (8) Solve the ODE
$$
    u'' + u' + \frac{5}{4} u = 0
$$
with initial conditions $u(0) = 3$ and $u'(0) = 1$.  Plot the solution and comment on its behavior as $t \rightarrow \infty$.

<!-- grade_id: A3-b1, points: 6, grade: true, solution: true -->
#
# YOUR CODE HERE



<!-- grade_id: A3-b2, points: 2, grade: true, solution: true -->
```python
#
# YOUR CODE HERE


```

## Question 4

(5) Plot something fun making sure to label the axes and colorbar if appropriate.  Use the [matplotlib gallery](http://matplotlib.org/gallery.html) for inspiration.

<!-- grade_id: A4, points: 5, grade: true, solution: true -->
```python
#
# YOUR CODE HERE


```