```python
%matplotlib inline
%precision 16
from __future__ import print_function
import numpy as np
import matplotlib.pyplot as plt
```

Before you turn this problem in, make sure everything runs as expected. First, restart the kernel (in the menubar, select Kernel $\rightarrow$ Restart) and then run all cells (in the menubar, select Cell $\rightarrow$ Run All).

Make sure you fill in any place that says YOUR CODE HERE or "YOUR ANSWER HERE", as well as your name and collaborators below:

# Homework 0:  Introduction

## Question 1

Compute the solution to the following sets of equations.  **FULLY** justify your solution (do not just write the answer).

If you want to review some of these concepts check out Strang's [Linear Algebra](https://clio.columbia.edu/catalog/10612929) text or for only the fundamentals check out Strang's [The Fundamental Theorem of Linear Algebra](http://www.jstor.org/stable/2324660?seq=1#page_scan_tab_contents)

**(a)** (5) Solve $A x = b$ where
$$
    A = \begin{bmatrix}
        2 & 1 \\
        1 & 3
    \end{bmatrix} ~~~~ 
    b = \begin{bmatrix}
        1 \\
        -2
    \end{bmatrix}
$$

<!-- grade_id: A1-a, points: 5, grade: true, solution: true -->
----
#Renjie Li, UNI = RL2932

<h1>Solution 1(a):</h1>

<h3>Method I :</h3>
With the concept of elimination, the above linear equation can be treated as two equations with two unknowns 
($x_{1}$and $x_{2}$are two components of vector $\bf{x}$):
  \begin{align}
  2x_{1}+x_{2}&=1  \\ 
  x_{1}+3x_{2}&=-2  
  \end{align}
   Substract 1/2 times equation 1 from equation 2:
  \begin{align}
  2x_{1}+x_{2}&=1 \\
  2.5x_{2}&=-2.5
  \end{align}
  Solution:
  $$
  x_{2}=\frac{-2.5}{2.5} = -1 \\
  x_{1}=\frac{1-x_{2}}{2}= 1
  $$
  Thus:
  $$
  x = \begin{bmatrix}
  x_{1} \\
  x_{2}
  \end{bmatrix}
  = \begin{bmatrix}
  1 \\
  -1
  \end{bmatrix}
  $$
  



<h3>Method II :</h3>
With the concept of inverse matrices, we need to find the inverse matrix of $A$ :
  
  Since the determinant of $A$ is $ 2\cdot 3-1\cdot 1=5$ , $A$ matrix is invertible .
  And its inverse matrix is :
  $$
  A^{-1}=\begin{bmatrix}
  2 & 1\\
  1 & 3
  \end{bmatrix}^{-1}
  = \frac{1}{5}\begin{bmatrix}
  3 & -1\\
  -1 & 2
  \end{bmatrix} 
  = \begin{bmatrix}
  \frac{3}{5} & -\frac{1}{5} \\
  -\frac{1}{5} & \frac{2}{5}
  \end{bmatrix}
  $$
  Thus:
  $$
  x = A^{-1}Ax = A^{-1}b \\
  x =\begin{bmatrix}
  \frac{3}{5} & -\frac{1}{5} \\
  -\frac{1}{5} & \frac{2}{5}
  \end{bmatrix}
  \begin{bmatrix}
  1 \\
  -2
  \end{bmatrix}
  =\begin{bmatrix}
  1 \\
  -1
  \end{bmatrix}
  $$
  
  
----  

**(b)** (5) Solve the system of equations:
\begin{align}
    2x + 3y &= 1 \\
    6x + 9y &= 3
\end{align}

<!-- grade_id: A1-b, points: 5, grade: true, solution: true -->
----
<h1>Solution 1(b):</h1>

With the concept of elimination, substract 3 times equation 1 from equation 2 :
$$
0y=0
$$
Every $y$ here satisfies the above equation. 
In fact there is only one equation, since the two equations in question description represent the same line.
We can choose the value of $y$ freely and then determine x by $x = \frac{1-3y}{2} $

----

**(c)** (5) Why will I not be able to solve $Ax=b$ for
$$
    A = \begin{bmatrix}
        2 & 1 \\
        1 & 3 \\
        0 & 4
    \end{bmatrix} ~~~~ 
    b = \begin{bmatrix}
        2 \\
        0 \\
        1
    \end{bmatrix}?
$$
Describe why this is a problem in terms of the column, row, left-null, and null spaces of $A$ and how that relates to the vector $b$.

<!-- grade_id: A1-c, points: 5, grade: true, solution: true -->
----
<h1>Description 1(c):</h1>

$Ax=b$ is solvable only if $b$ is in the column space of $A$.
$A$ is a 3 by 2 matrix, 2 columns cannot fill 3-dimensional vector space , which means some $b$ are not in the column space of $A$.

The row echelon form of $A$ is :
$ \begin{bmatrix}
2 & 1\\
0 & 2.5\\
0 & 0
\end{bmatrix}
$

It has two pivots, which means both two columns will define the column space of $A$.Thus the column space of $A$ is :
$
 x_{1} \begin{bmatrix}
 2 \\ 1 \\ 0
 \end{bmatrix}
 + x_{2}\begin{bmatrix}
 1 \\ 3 \\ 4
 \end{bmatrix}
$

The given vector $b$ here is not in the column space of $A$. That is why this equation is not solvable.


By the way, we can also use elimination on augmented matrix:
$
\begin {bmatrix}
2 & 1 & b_{1}\\
1 & 3 & b_{2}\\
0 & 4 & b_{3}
\end{bmatrix}
\rightarrow
\begin {bmatrix}
2 & 1 & b_{1} \\
0 & 2.5 & b_{2}-0.5b_{1} \\
0 & 0 & b_{3}-1.6b_{2}+0.8b_{1} 
\end{bmatrix}
$

$Ax=b$ is solvable when every entry in last row equals zero. In this case, $1-1.6\cdot 0+0.8\cdot 2 = 2.6 \neq 0 $ .
We can not solve this equation.

----

## Question 2

**(a)** (10) Write a function that computes
$$
    \sum^\infty_{n=1} \frac{a^n}{b^{n-1}}
$$
until the difference between subsequent partial sums is less than the given tolerance $T$.  Return the computed sum.  Make sure to include a way for the function to exit if the partial sums do not satisfy the above criteria (the sum may not be convergent for instance).

<!-- grade_id: A2, points: 0, grade: false, solution: true -->
```python
def compute_sum (a,b,Tol):
    if b == 0 :  # zero is not allowed to be denominator
        raise ValueError
    sum = 0
    n = 1
    partial_sum = (a ** n) / (b**(n-1))
    if abs(a/b) > 1:  # the sum may not be convergent if a/b >1 (difference between subsequent partial sums > Tol)
        raise ValueError
    bool = True
    while bool :
        if (abs(partial_sum)) > Tol :
            sum += partial_sum
            n += 1
            partial_sum = (a ** n) / (b**(n-1))
        else:
            bool = False
    return sum
    
compute_sum(-2.0,3.0,1e-16)
```

Output:
```
-1.1999999999999997
```

<!-- grade_id: T2-a, points: 10, grade: true, solution: false -->
```python
import numpy
numpy.testing.assert_allclose(compute_sum(-2.0, 3.0, 1e-16), -1.2)
numpy.testing.assert_allclose(compute_sum(1.0, 2.0, 1e-16), 2.0)
try:
    compute_sum(-2.0, 1.0, 1e-16)
except ValueError:
    pass
else:
    assert False

```

```python
import numpy
numpy.testing.assert_allclose(compute_sum(-2.0, 3.0, 1e-16), -1.2)
numpy.testing.assert_allclose(compute_sum(1.0, 2.0, 1e-16), 2.0)
try:
    compute_sum(-2.0, 3.0, 1e-16)
except ValueError:
    pass
else:
    assert False
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-26-fb61ffa9d662> in <module>()
      7     pass
      8 else:
----> 9     assert False

AssertionError:
```

**(b)** (5) Explore different tolerances for your function above for $a=-2$ and $b=3$.  Plot the difference
$$
    \left| ~ \sum^\infty_{n=1} \frac{a^n}{b^{n-1}} - (-1.2) ~\right |
$$
versus the value of the tolerance.  The plotting command `loglog` may be useful to effectively visualize the problem.  What do you observe?

Hint:  Try using tolerances in the range $T \in [10^{-30}, 10^{-2}]$.

<!-- grade_id: A2-b1, points: 4, grade: true, solution: true -->
```python
T=np.logspace(-30,-2,num=200,base=10)
Diff=[]
Tolerance=[]
for i in T:
    dif = abs(compute_sum(-2,3,i)+1.2)
    Tolerance.append(i)
    Diff.append(dif)
    
plt.loglog(Tolerance,Diff,'g',basex=10,basey=10)
# As the value of tolerance raises, the value of difference also becomes bigger

```

## Question 3

**(a)** (7) Compute the first 3 terms of the Taylor series of the function
$$
    f(x) = e^{-x^2} \sin( x - \pi)
$$
centered at $x_0 = \pi$.

<!-- grade_id: A3-a, points: 7, grade: true, solution: true -->
----
<h1>Solution 3(a):</h1>


Derivatives:
$$\begin{aligned}
f'(x) &= e^{-x^{2}}(2xsin(x)-cos(x)) \\
f''(x) &= e^{-x^{2}}[(3-4x^{2})sin(x)+4xcos(x)]
\end{aligned}
$$

Taylor polynomials:
$$T_N(x) = \sum^N_{n=0} \frac{f^{(n)}(x_0)\cdot(x-x_0)^n}{n!} $$
For centered at $x_{0}=\pi $
$$T_N(x) = \sum^N_{n=0} \frac{f^{(n)}(\pi)\cdot(x-\pi)^n}{n!}\Rightarrow $$

$$
\begin{aligned}
T_2(x) & = e^{-\pi^{2}}sin(\pi - \pi)+ e^{-\pi^{2}}[2\pi sin(\pi)-cos(\pi)](x-\pi)+\frac{1}{2}e^{-\pi^{2}}[(3-4\pi^{2})sin(\pi)+4\pi cos(\pi)](x-\pi)^{2}  \\
& = \pi e^{-\pi^{2}}(x-\pi)-2e^{-\pi^{2}}(x-\pi)^{2}
\end{aligned}
$$

----

**(b)** (8) Solve the ODE
$$
    u'' + u' + \frac{5}{4} u = 0
$$
with initial conditions $u(0) = 3$ and $u'(0) = 1$.  Plot the solution and comment on its behavior as $t \rightarrow \infty$.

<!-- grade_id: A3-b1, points: 6, grade: true, solution: true -->
----
<h1>Solution 3(b): </h1>

<h3>Method I :</h3>

The characteristic equation of the above ODE is :
$$
r^{2}+r+\frac{5}{4}=0
$$
whose roots are $r_{1,2}=-\frac{1}{2}\pm 1i $.
The general solution shoule be:
$$
\begin{aligned}
u(t) &= C_{1}e^{- \frac{1}{2}t}cos(t)+C_{2}e^{- \frac{1}{2}t}sin(t) \\
u'(t) &= -\frac{C_{1}}{2}e^{- \frac{1}{2}t}cos(t) - C_{1}e^{- \frac{1}{2}t}sin(t) - \frac{C_{2}}{2}e^{- \frac{1}{2}t}sin(t)+C_{2}e^{- \frac{1}{2}t}cos(t)
\end{aligned}
$$
Apply $u(0)=3$ and $u'(0)=1$ into the above two equations, we can get the values of $C_{1}$ and $C_{2}$ :
$$
C_{1}=3 \\
C_{2}=\frac{5}{2} 
$$
Then:
$$
u(t)=3e^{- \frac{1}{2}t}cos(t)+\frac{5}{2}e^{- \frac{1}{2}t}sin(t)
$$



<h3>Method II: </h3>

We can solve it with Laplace Transform:
$$
\mathcal{L}{(u''(t))}+\mathcal{L}{(u'(t))}+\frac{5}{4}\mathcal{L}{(u(t))}=0 \Rightarrow \\
s^{2}U(s)-su(0)-u'(0)+sU(s)-f(0)+\frac{5}{4}U(s)=0
$$
Apply $u(0)=3$ and $u'(0)=1$ into the above equations:
$$
s^{2}U(s)-3s-1+sU(s)-3+\frac{5}{4}U(s)=0 \Rightarrow \\
U(s) = \frac{3s+4}{s^{2}+s+\frac{5}{4}}
= 3\frac{(s+\frac{1}{2})}{(s+\frac{1}{2})^{2}+1^{2}}+\frac{5}{2} \frac{1}{(s+\frac{1}{2})^{2}+1^{2}}
$$
Then take the inverse transform we can get:
$$
u(t)=3e^{- \frac{1}{2}t}cos(t)+\frac{5}{2}e^{- \frac{1}{2}t}sin(t)
$$

----

<!-- grade_id: A3-b2, points: 2, grade: true, solution: true -->
```python
import numpy as np
t = numpy.linspace(0, 50, 100)
u = 3*np.exp(-0.5*t)*np.cos(t)+2.5*np.exp(-0.5*t)*np.sin(t)
plt.plot(t, u,'r')
plt.xlabel("t")
plt.ylabel("u(t)")
plt.title("Question 3(b)")
#plt.ylim([-0.01,0.01]) //Uncomment these two lines will figure out 
#plt.xlim([20,50])      //the value of u(t) is almost zero when t -> infinity
```

----
As $t \rightarrow \infty$, $e^{- \frac{1}{2}t} \rightarrow 0$. 

Thus $u(t)=3e^{- \frac{1}{2}t}cos(t)+\frac{5}{2}e^{- \frac{1}{2}t}sin(t) \rightarrow 0$

----

## Question 4

(5) Plot something fun making sure to label the axes and colorbar if appropriate.  Use the [matplotlib gallery](http://matplotlib.org/gallery.html) for inspiration.

<!-- grade_id: A4, points: 5, grade: true, solution: true -->
```python

import numpy
import matplotlib
y=numpy.linspace(0,15,100)
x=0.2*numpy.sin(y)
x2=0.2*numpy.cos(y)
x3=0.1+0.2*numpy.sin(0.8*y)
plt.plot(x,y,'r')
plt.plot(x2,y-1.5,'r')
xline = [-0.18, 0.18,-0.18 ,0.18, -0.18, 0.18,-0.1,0.1,-0.1,0.1,-0.1,0.1,-0.1,0.1,-0.1,0.1,-0.1,0.1]
yline = [1.8, 1.8, 4.8, 4.8, 7.8, 7.8,1,1,4,4,7,7,2.5,2.5,5.5,5.5,8.5,8.5]
pair_x_array = np.reshape(xline, (-1, 2))
pair_y_array = np.reshape(yline, (-1, 2))
for i, pair_x in enumerate(pair_x_array):
    pair_y = pair_y_array[i]
    plt.plot(pair_x, pair_y, 'g', linewidth=2)
plt.xlim([-1,1])
plt.ylim([0.5,9])
plt.title("DNA Strand in 2D")
```

```python

```